---
title: 'Python Coaching Classes in Pune at SourceKode Institute'
description: 'A good Python Class is essential to earn a handsome salary and choosing the Best Python Classes & Training Course is highly crucial towards getting a Great Job as a Python Developer. Learn concepts of python like Django python which is famous nowadays. Attend a Demo Session with sourceKode Training Institute and Choose the best. We provide 100% assured placements.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
